<h1> Matemática Discreta</h1>

<h4> Grao en Enxeñaría Informática. &nbsp; &nbsp; Curso 2021-2022 &nbsp; &nbsp; &nbsp; (<i>2ª oportunidade</i>)</h4>



## Breve explicación da proba

A presente folla non é **.sagews** (*sagemath woksheet*), senón **.ipynb** (*jupyter notebook*). Coma nunha sage worksheet, se precisas abrir un grupo de execución, simplemente toca co rato na barra que separa dous grupos. Dado que os comandos corresponde a sagemath, para executar hai que pulsar shift + return, coma sempre.

Nesta folla vaste atopar con celdas de comentarios (coma esta), celdas con funcións definidas por min (que non podes modificar) que debes executar para poder usar a función. Despois, hai celdas nas que se che pide definir unha función que veñen seguidas dunha celda con casos de test que deberá satisfacer a túa definición.

A modo de exemplo, fíxate nas dúas celdas que seguen:

```
def restar(a,b):  
    """
    devolve a resta dos números a e b   (Unha explicación do que se quere coa función)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```
# algúns tests para comprobar que restar funciona como se espera
from nose.tools import assert_equal
assert_equal(restar(1,2), -1)
assert_equal(restar(0,-5), 5)
```

Na primeira terías que definir a función `restar` con parámetros \(a\) e \(b\) que ten que devolver \(a - b\). 

Tes que escribir o teu código no lugar indicado, eliminando a liña

`raise NotImplementedError()`.

Na segunda verás os tests que deberá pasar a túa función, `restar` [^1]

- `restar(1,2) == -1`
- `restar(0,-5) == 5`
 

Na parte superior de cada celda de tests está a puntuación asignada a pregunta.

[^1] *Podes abrir unha celda e facer ti outras comprobacións adicionais.*



# Exame



Lembrarás que, aparte da orde, a diferencia entre unha **lista** e un **conxunto** é que as listas poden conter elementos repetidos, e os conxuntos non. En `SageMath` a partir dunha lista $L$ podemos definir un conxunto facendo _Set_\($L$\), ```Set(L)``` pero o resultado non será unha lista, senón un conxunto. Pretendemos facer algo similar pero que dea como resultado unha lista, **sen utilizar** *Set*.

Podes utilizar as funcións e métodos que aprendimos para o manexo de listas,  como  [$x+3$ **for** $x$ **in** $L$], $L$ **+** $M$, [$x$  **for** $x$ **in** $L$ **if** $x >3$ ], $L$**.append( )**, **len(**$L$**)**, **sorted(**$L$**)**, etc.

## Pregunta 1. &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; (2 ptos.)

**Definir** a función **S** para eliminar as repeticións e ordenar unha lista de números. É dicir, collida unha lista $L$, **S(**$L$**)** será a lista formada por todos os elementos que aparecen en $L$, pero sen repeticións e ordenada. Por exemplo, **S(**$[3,1,2,1]$**)** = $[1,2,3]$

<p><hr width="50%" color="green" size="10px" /></p>
 <h6> Lembra:
<ul>
        <li>Empeza a escribir <strong>debaixo</strong> de onde di &nbsp; # YOUR CODE HERE,</li>
        <li><font color="red">Borra a liña&nbsp; &nbsp;<strong> raise NotImplementedError()</strong></font>,</li>
        <li> <strong>Non</strong> repitas def S(L) ...  nin cambies o nome da función.</li>
    <li> Non utilices <i>Set</i></li>
    </ul> 
</h6>



<!-- grade_id: answ_1, points: 0, grade: false, solution: true -->
```sagemath
def S(L):
    """
    Devolve unha lista cos elementos de L sen repeticións e en orden.
    Lembra eliminar a liña raise NotImplementedError()
    """
    # YOUR CODE HERE
    sol =[]
    for x in L:
        if not x in sol:
            sol.append(x)
    return sorted(sol)
```

```sagemath
L.sort()
    sol = []
    sol.append(L[0])
    i = 1
    for i in [0..len(L)-1]:
        m = sol.count(L[i])
        if m < 1:
            sol.append(L[i])
    return sol
```

Output:
```
File "<ipython-input-276-6f0e609ec436>", line 2
    sol = []
    ^
IndentationError: unexpected indent
```

```sagemath
sol =[]
for x in L:
    if not x in sol:
        sol.append(x)
    sorted(sol)
```

```sagemath
S([])
```

Output:
```
[]
```

```sagemath
S([3,1,2,1,2,3])
```

Output:
```
[1, 2, 3]
```

<!-- grade_id: test_1, points: 2, grade: true, solution: false -->
```sagemath
# algúns tests que debe pasar a túa lista
from nose.tools import assert_equal
assert_equal(S([2,1]), S([1,2]))
assert_equal(S([3,1,2,1]), [1,2,3])
assert_equal(S([3,1,2,1]), [1,2,3])
assert_equal( S([[3,1,2,1]]), [[3,1,2,1]])

# para comprobar que non usas a función Set
from unittest.mock import patch
with patch('__main__.Set') as mock_Set:
    S([1,2])
mock_Set.assert_not_called()

```

## Pregunta 2. &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; (1 pto.)
**Utilizar** a función **S** para definir a función **U** que fai  a *unión* de dúas listas. É dicir, collidas dúas listas, $L$ e $M$, **U(**$L,M$**)** será a lista formada por todos os elementos que aparecen en $L$ ou en $M$, pero sen repeticións e en orden. Por exemplo, **U(**$[1,2,1,3],[2,4,3]$**)** = $[1,2,3,4]$

<!-- grade_id: answ_2, points: 0, grade: false, solution: true -->
```sagemath
def U(L,M):
    """
    Devolve os elementos en L ou en M, sen repeticións e en orden.
    """
    # YOUR CODE HERE
    u = L + M
    sol = S(u)
    return sol
```

```sagemath
U([1,2,1,3], [2,4,3])
```

Output:
```
[1, 2, 3, 4]
```

<!-- grade_id: test_2, points: 1, grade: true, solution: false -->
```sagemath
# Algúns exemplos para probar a túa función:
from nose.tools import assert_equal
assert_equal(U([1,2,1,3],[2,4,3]), [1,2,3,4])
assert_equal(U([2,1,2],[]), [1,2])

# test para verificar que se usa a función pedida
# eliminamos S temporalmente
old_S = S
del S

try:
    U(L,M)
# se dá NameError, é que usa S
except NameError:
    pass
# se non o dá, non a usa
else:
    raise AssertionError("U non usa S")
# restauramos a función
finally:
    S = old_S
    del old_S
```

## Pregunta 3. &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; (2 ptos.)

**Definir** a función **P** para determinar todas as *partes* dunha lista. É dicir, se $L$ é unha lista de números, **P(**$L$**)** será a lista ordenada de todas as posibles listas ordenadas que se poden facer con elementos de $L$ distintos. Por exemplo, **P(**$[1,2,1]$**)** = $[[\,\,], [1], [1, 2], [2]]$

<!-- grade_id: ans_3, points: 0, grade: false, solution: true -->
```sagemath
def P(L):
    """
    Sendo L unha lista de números,
    devolve a lista ordenada de todas as sublistas de L con elementos distintos e en orden
    """
    # YOUR CODE HERE
    l = []
    sol =  []
    m = S(L)
    lista = Combinations(m)
    return lista.list()
```

```sagemath
P([])
```

Output:
```
[[]]
```

```sagemath
P([2,3,3,4])
```

Output:
```
[[], [2], [3], [4], [2, 3], [2, 4], [3, 4], [2, 3, 4]]
```

<!-- grade_id: test_3, points: 2, grade: true, solution: false -->
```sagemath
# Algúns tests que debe pasar a túa función
from nose.tools import assert_equal
assert_equal(P([1,2,1]), [[ ], [1], [2],[1,2]])
assert_equal(P([2,1,1]),P([1,2]))
assert_equal(P([]), [[ ]])

```

## Pregunta 4. &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; (2 ptos.)

Imaxina que para camiñar podes dar pasos de 1 metro ou saltos de 2 metros.

Coa cadea **'12112'** denotamos unha posible forma de camiñar 7 metros: primerio un paso de $1\, m$, despois un salto de $2\, m$, logo dous pasos de $1\, m$ e finalmente un salto de $2\, m.$

**Definir** a función **ForCam** que toma como parámetro un número $n$ e devolve todas as formas posibles de camiñar $n$ metros con pasos de $1\, m$ e/ou saltos de $2\, m.$

Lembra que, por exemplo, a cadea '121' podemos obtela como ''.join(['1','2','1'])



<!-- grade_id: answ_4, points: 0, grade: false, solution: true -->
```sagemath
def ForCam(n): 
    """
    Devolve unha lista con todas as formas de camiñar n metros con pasos de 1 metro ou saltos de 2 metros
    """
    res = []
    for i in range(floor(n/2)+1):
        res += [''.join(p) for p in Permutations('2'*i+'1'*(n-2*i))]
    return res
```

```sagemath
ForCam(4)
```

Output:
```
['1111', '211', '121', '112', '22']
```

<!-- grade_id: test_4, points: 2, grade: true, solution: false -->
```sagemath
# algúns tests para a túa función
assert ForCam(2) == ['11', '2']

```

## Algoritmos recursivos           \(3 ptos.\)




### Pregunta 5.a\)           \(1.5 ptos.\)

<p>
Definir unha función <strong>num_cp</strong> mediante un <font color="red"><strong>algoritmo recursivo</strong></font> que para $n \ge 0$ nos devolve o numero de cadeas de $n$ bits cunha cantidade par de ceros.
</p>



<!-- grade_id: answ_5a, points: 1.5, grade: true, solution: true -->
```sagemath
@cached_function
def num_cp(n):
    """
    Sendo n un número enteiro non negativo,
    devolve o número de cadeas de n bits
    cun número par de ceros.
    Lembra que se pide un ALGORITMO RECURSIVO
    """
    if n <= 1: return []
    else:
        comb = [''.join(p) for p in Permutations('01'*(n-2),n-2)]
        sol = ['00'+c for c in comb] + ['10'+c for c in ceros_pares(n-2)] + ['1'+c for c in ceros_pares(n-1)]
    return len(sol)
```

### Pregunta 5.b) &nbsp; &nbsp; &nbsp; &nbsp; &nbsp; (1.5 ptos.)

<p>
    Definir unha función <strong>ceros_pares</strong> mediante un <font color="red"><strong>algoritmo recursivo</strong></font> para que para $n \ge 0$ nos devolve todas as cadeas de $n$ bits cunha cantidade par de ceros.
</p>



<!-- grade_id: answ_5b, points: 0, grade: false, solution: true -->
```sagemath
@cached_function
def ceros_pares(n):
    """
    Sendo n un número enteiro non negativo,
    devolve a lista de cadeas de n bits
    cun número par de ceros.
    Lembra que se pide un ALGORITMO RECURSIVO
    """
    # YOUR CODE HERE
    if n <= 1: return []
    else:
        comb = [''.join(p) for p in Permutations('01'*(n-2),n-2)]
    return ['00'+c for c in comb] + ['10'+c for c in ceros_pares(n-2)] + ['1'+c for c in ceros_pares(n-1)]
```

```sagemath
len(ceros_pares(5))
```

Output:
```
19
```

<!-- grade_id: test_5b, points: 1.5, grade: true, solution: false -->
```sagemath
### Algúns tests para probar a túa función
from nose.tools import assert_equal, assert_in, assert_not_in
assert_not_in([0,0], ceros_pares(2))
assert_in('001', ceros_pares(3))

```

```sagemath

```